# Introduction to Python and Natural Language Technologies

__Laboratory 01, Introduction to Python__

__Feb 11, 2021__

__Submission deadline: Feb 14 23:59__

__Submissions must be made through Github Classroom__

You are required to complete all tasks correctly (or almost correctly) until == PASSING LEVEL ==. <br />
In order to receive a point for a projected grade, you have to complete the exercises until == EXTRA LEVEL ==.<br />
There are additional exercises at the end of the notebook.
We recommend completing these as well but they are not required for the Extra Level.
You can make up for mistakes by solving more exercises than the ones until Passing/Extra level.

You are free to use any material (lecture, Stackoverflow etc.) except full solutions but you should first try to figure them out on your own.

We use an [nbgrader](https://nbgrader.readthedocs.io/en/stable/index.html) for partially autograding the assignments. Please **DO NOT** delete or copy cells. You can add new cells but your solutions should go into the designated cells.

# `range()` practice

## Print the numbers between 60 and 69 inclusive.

<!-- grade_id: cell-306d37c936005bde, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Print the first 10 square numbers starting from 1.

<!-- grade_id: cell-b5c9f9ae3534a9b4, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Print the numbers from 29 to 20 (backwards).

<!-- grade_id: cell-3b0aefde40906e95, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Reading user input

User input can be read with the `input()` function. It always returns a string which can be converted to integer using the `int()` function.

<!-- grade_id: cell-3acd7027f4f934a8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Read a number N and print the numbers from 1 to N.

<!-- grade_id: cell-0a36206459fb1dd0, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Read a number N and print the first N square numbers starting from 1.

<!-- grade_id: cell-834b3f1952f3a615, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Read numbers until 0 and print their sum.

<!-- grade_id: cell-445648d32d55a5e5, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Read numbers until 0 and print their mean and standard deviation without using a list.

Do not include 0 in the statistics.

<!-- grade_id: cell-25bc5f05e52e11c5, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Lists

## Print the odd numbers from this list.

```python
l = [2, 3, -2, -7, 0, 2, 3]
```

<!-- grade_id: cell-f0e76a9ae83db0ba, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Read numbers until 0 and collect them in a list. Print the elements backwards (zero not included).

For example:

```
input: 3 4 2 12 3 0
output: 3 12 3 4 3 ```

<!-- grade_id: cell-89c5428a352b0b90, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Read a number N and then read exactly N numbers. Print them backwards.

For example:

```
input: 4 -1 0 2 12
output: 12 2 0 -1
```

<!-- grade_id: cell-28fe536393b9cbff, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Functions

The following exercises contain function skeletons. Finish these functions.
If you work correctly, the assert statements should not fail, however, passing all tests does not guarantee that your solution is correct.
You are free to add more tests.

## Write a function that takes an integer parameter (N) and returns a list of the first N prime numbers starting from 2.

<!-- grade_id: cell-32bfa66ad55808d0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1338b4225b785617, points: 0, grade: true, solution: false -->
```python
assert(get_n_primes(3) == [2, 3, 5])
assert(get_n_primes(0) == [])
```

# ================= PASSING LEVEL ====================

## Write a function that takes an integer parameter (N) and returns the Nth Fibonacci number.

The Fibonacci series is defined as:

$$
F_n = F_{n-1} + F_{n-2} \text{ for } n>1\\
F_0 = 0\\
F_1 = 1
$$

The first few elements of the Fibonacci series are 0, 1, 1, 2, 3, 5, 8, 13 etc.

Raise a `ValueError` if $N<0$.

<!-- grade_id: cell-bed163595b1306c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c956085d6aa9b32d, points: 0, grade: true, solution: false -->
```python
assert(get_nth_fibonacci(0) == 0)
assert(get_nth_fibonacci(1) == 1)
assert(get_nth_fibonacci(2) == 1)
assert(get_nth_fibonacci(5) == 5)
assert(get_nth_fibonacci(7) == 13)

try:
    get_nth_fibonacci(-2)
except ValueError:
    pass
else:
    AssertionError("TEST FAILS. Negative index should raise a ValueError.")
```

## Write a function that takes an integer parameter (N) and returns a list of the first N Fibonacci numbers starting from $F_1$ to $F_N$ inclusive.

You are encouraged to reuse previous functions.

<!-- grade_id: cell-a483efba1cf84902, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-70487d81d741784a, points: 0, grade: true, solution: false -->
```python
assert(get_n_fibonacci(4) == [1, 1, 2, 3])
assert(get_n_fibonacci(8) == [1, 1, 2, 3, 5, 8, 13, 21])
assert(get_n_fibonacci(0) == [])
```

## Write a function that takes two integers, A and B and returns B-A Fibonacci numbers starting from the Ath Fibonacci number to the (B-1)th number. See the example below.

<!-- grade_id: cell-bdda0121e9b4384e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-009e90c9d285f80e, points: 0, grade: true, solution: false -->
```python
assert(get_range_fibonacci(2, 5) == [1, 2, 3])
assert(get_range_fibonacci(5, 5) == [])
assert(get_range_fibonacci(1, 5) == [1, 1, 2, 3])
```

# Default and keyword arguments

## Write a function that computes the sum of a list with an optional starting value. If no starting value is provided, it should sum from 0.

<!-- grade_id: cell-33aba3b44cbe7d79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-11357ac08ed33343, points: 0, grade: true, solution: false -->
```python
assert(sum_list([1, 2, 3]) == 6)
assert(sum_list([1, 2, 3], 5) == 11)
assert(sum_list([1, 2, 3, 5, 11, 3213], 5) == 3240)
```

## Write a function that takes two numbers and an arithmetic operator as a string ("+", "-", "\*" or "/") and returns the result of the arithmetic operation on the two numbers. The default operation should be addition.

<!-- grade_id: cell-f2dad9c42958066e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7ed54437386acf59, points: 0, grade: true, solution: false -->
```python
assert(arithmetic(2, 3) == 5)
assert(arithmetic(2, 3, "+") == 5)
assert(arithmetic(2, 3, "-") == -1)
assert(arithmetic(2, 3, "*") == 6)
assert(arithmetic(2, 3, "/") == 2/3)
```

## Write a function that takes three parameters. The third is a callable (function) that takes two parameters. The function should call its third parameter with the first two as parameters. If the third parameter is not specified the function should add the two arguments.

This is very similar to the previous function but the third parameter is a function instead of a string.

<!-- grade_id: cell-3d76a413c7125655, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c77dfa723cea60a1, points: 0, grade: true, solution: false -->
```python
assert(call_func(3, 4, product) == 12)
assert(call_func("foo", "bar") == "foobar")
```

# ================= EXTRA LEVEL ====================

# Additional exercises

## Create a function that take a list and a predicate (function with boolean return value) as parameters and returns a new list of those elements which the predicate return True.

A predicate is a function that takes one element and return `True` or `False`, for example `is_even`, `is_prime`.

If you implemented The following tests should run.

<!-- grade_id: cell-244a0447cef16213, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-07ddbdcd2afcdf4c, points: 0, grade: true, solution: false -->
```python
l1 = [1, 2, 3, 4, 19, 35, 11]
l2 = [1, 4, 9, 16, 25]
assert(filter_list(l1, is_odd) == [1, 3, 19, 35, 11])
assert(filter_list(l1, is_prime) == [2, 3, 19, 11])
assert(filter_list(l2, is_prime) == [])
```

## Reduce is a function that applies a two argument function against an accumulator and each element in the sequence (from left to right) to reduce it to a single value. If no initial value is provided, the accumulator is initialized with the return value of the function run on the first two elements of the sequence.

```
reduce([1, 2, 3], product) ---> 6
reduce([1, 2, 3], product, accumulator=10) ---> 60
reduce(["foo", "bar"], string_addition) ---> "foobar"
reduce(["foo", "bar"], string_addition, accumulator="hello") ---> "hellofoobar"
```

<!-- grade_id: cell-6fb5723159b26749, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c43e405f72ada059, points: 0, grade: true, solution: false -->
```python
l1 = [1, 2, -1, 5]
l2 = ["foo", "bar", "hello"]

assert(reduce(l1, add) == 7)
assert(reduce(l1, add, 10) == 17)
assert(reduce(l2, string_len_add, 0) == 11)
```

## Use your reduce function for the following operations:

1. count the number of odd elements in a list of integers,
2. find the maximum of a list of integers,
3. find the longest string in list of strings.

Test your solutions.

<!-- grade_id: cell-1554306956c531d6, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-959dbab0fb50d84d, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Implement qsort. Qsort sorts a list in place using the quicksort algorithm.

<!-- grade_id: cell-a9a76c9ca712b9ee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-034cb73519ff059b, points: 0, grade: true, solution: false -->
```python
l1 = [10, -1, 2, 11, 0]
l2 = [5, 4, 3, 1, 2]
qsort(l1)
qsort(l2)
assert(l1 == [-1, 0, 2, 10, 11])
assert(l2 == [1, 2, 3, 4, 5])
```

## Implement bubble sort (in place).

<!-- grade_id: cell-40c165b7b3725cc9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1de3197545574d72, points: 0, grade: true, solution: false -->
```python
l1 = [10, -1, 2, 11, 0]
l2 = [5, 4, 3, 1, 2]
bubblesort(l1)
bubblesort(l2)
assert(l1 == [-1, 0, 2, 10, 11])
assert(l2 == [1, 2, 3, 4, 5])
```