# Introduction to Python and Natural Language Technologies

__Laboratory 02, Type system and built-in types__

__Feb 18, 2021__


We will use the wikipedia module for some exercises. It can be installed by running pip from Jupyter:

The Python Standard Library has solutions for many of these tasks but you should solve them without relying on the standard library.

```python
! pip install wikipedia
```

# 1. Define a function that takes a sequence as its input and returns whether the sequence is symmetric. A sequence is symmetric if it is equal to its reverse.

<!-- grade_id: cell-9063e2fba9c7157d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-58b3ad7e361196e7, points: 0, grade: true, solution: false -->
```python
assert is_symmetric([1]) == True
assert is_symmetric([None, None]) == True
assert is_symmetric([]) == True
assert is_symmetric([1, 2, 3, 1]) == False
assert is_symmetric([1, "foo", "bar", "foo", 1]) == True
assert is_symmetric("abcba") == True
assert is_symmetric(list(range(100)) + list(range(99, -1, -1))) == True
assert is_symmetric(list(range(100)) + list(range(100, 0, -1))) == False
assert is_symmetric([1, "1"]) == False
```

# 2. Define a function that takes a sequence and counts the frequency of each element in the list. The function should return a dictionary of the counts.

Do not use the `collections` module.

<!-- grade_id: cell-d36bfd560f1b29a3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5175730489a10307, points: 0, grade: true, solution: false -->
```python
# test for list
assert count_elements([1, 2, 0, 1, 1]) == {1: 3, 2: 1, 0: 1}
# test for tuple
assert count_elements((1, 2, 0, 1, 1)) == {1: 3, 2: 1, 0: 1}
# test empty list
assert count_elements([]) == {}
# test mixed types
assert count_elements([1, "1", 1]) == {1: 2, "1": 1}
```

# 3. Define a function that takes a list of lists and decided if it is a matrix. A list of lists is a matrix if each nested list is the same length.

<!-- grade_id: cell-90b341f4b426d33d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-218d090c410748b5, points: 0, grade: true, solution: false -->
```python
assert is_matrix([]) == True
assert is_matrix([[], []]) == True
assert is_matrix([[1, 2], [2, 1]]) == True
assert is_matrix([[1], [2, 1]]) == False
assert is_matrix([[1, 2], [2, 1], [3, 4]]) == True
assert is_matrix([[1, 2], [2, 1], [3, 4, 5]]) == False
```

# 4. Define a function that takes a string as its input and return a dictionary with the character frequencies.

Do not use the `collections` module.

<!-- grade_id: cell-7fbaf8610be225a9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ecedb93a6c391539, points: 0, grade: true, solution: false -->
```python
assert char_freq("aba") == {"a": 2, "b": 1}
assert char_freq("") == {}
assert char_freq("ab" * 100 + "cd") == {"a": 100, "b": 100, "c": 1, "d": 1}
```

# 5. Define a function that computes word frequencies in a text.

You can split the text with `text.split()`:

<!-- grade_id: cell-4d4d473dd807f962, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1c299c2c935129e7, points: 0, grade: true, solution: false -->
```python
s = "the green tea and the black tea"
assert(word_freq(s) == {"the": 2, "tea": 2, "green": 1, "black": 1, "and": 1})
assert(word_freq("") == {})
```

# Wikipedia module

The following exercises use the `wikipedia` package. The basic usage is illustrated below.

The documentation is available [here](https://pypi.python.org/pypi/wikipedia/).

Searching for pages:

```python
import wikipedia

results = wikipedia.search("Budapest")
results
```

Downloading an article:

```python
article = wikipedia.page("Budapest")

article.summary[:100]
```

The content attribute contains the full text:

```python
type(article.content), len(article.content)
```

By default the module downloads the English Wikipedia. The language can be changed the following way:

```python
wikipedia.set_lang("fr")
```

```python
wikipedia.search("Budapest")
```

```python
fr_article = wikipedia.page("Budapest")
fr_article.summary[:100]
```

```python
wikipedia.set_lang("en")
```

# 6. Build a small Wikipedia corpus by downloading pages until their character count exceeds 100000. The corpus should be a list of strings.

You can download arbitrary pages but some pages redirect to disambiguation pages. You should not pick these. Make sure the same page is not included multiple times.

<!-- grade_id: cell-2e7be91b99158dca, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f26aa917928c62c8, points: 0, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus()
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 100000

# no duplicate pages
assert len(set(corpus)) == len(corpus)
```

# 7. Add a `min_char` parameter to `build_wikipedia_corpus` that defaults to 10000. You should stop building the corpus when this number is reached.

<!-- grade_id: cell-bf7ffeba0644e753, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-68c3de00731699f8, points: 0, grade: true, solution: false -->
```python
corpus = build_wikipedia_corpus(1000)
# check types
assert isinstance(corpus, list)
assert isinstance(corpus[0], str)

# check character count
sum_length = 0
for document in corpus:
    sum_length += len(document)
assert sum_length >= 1000
# shouldn't be too long
assert sum_length < 100000
```

# 8. Using your previously defined function compute the word frequencies of your Wikipedia corpus. Print the ten most frequent words and their frequency in the following format:

```
a<TAB>12
an<TAB>11
```

You can use str.split() for word tokenization:

```python
"This is a short sentence".split()
```

There are no tests for this task. Try to come up with your own tests.

<!-- grade_id: cell-8c167ceb2a4c65f1, points: 0, grade: true, solution: true -->
```python
# def get_corpus_word_frequencies(corpus):
# YOUR CODE HERE
raise NotImplementedError()
```

Your own tests go here:

<!-- grade_id: cell-d17a9e499c13bd35, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 9. Define a function that takes a matrix as an input represented as a list of lists. Return its transpose without changing the original matrix.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-adee24575dcbac8e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9b412cb3d0c0506a, points: 0, grade: true, solution: false -->
```python
m1 = [[1, 2, 3], [4, 5, 6]]
m2 = [[1, 4], [2, 5], [3, 6]]

assert transpose([[1]]) == [[1]]
assert transpose(m1) == m2
# m1 is unchanged
assert m1 == [[1, 2, 3], [4, 5, 6]]

# transpose is an involution (self-inverse)
assert transpose(transpose(m1)) == m1
```

# =================== PASSING LEVEL ===================

# 10. Define a function that takes a matrix as its input and returns the row-wise maximum elements as a list.

Raise a `ValueError` if the input is not a valid matrix. Don't forget to reuse your code.

<!-- grade_id: cell-21f91aada57c0332, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eb440f91ea977032, points: 0, grade: true, solution: false -->
```python
assert rowwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [3, 4, 5]

assert rowwise_max([list(range(1000))]) == [999]
assert rowwise_max(transpose([list(range(1000))])) == list(range(1000))
```

# 11. Define a function that takes a matrix as its input and returns the column-wise maximum elements as a list.

Try to reuse your previous functions.

<!-- grade_id: cell-bde0cb9cf8d75c89, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-142e6034ed15005b, points: 0, grade: true, solution: false -->
```python
assert columnwise_max(
    [
        [1, 2, 3],
        [1, 4, 3],
        [1, 5, 3],
    ]
) == [1, 5, 3]

assert columnwise_max([list(range(1000))]) == list(range(1000))
assert columnwise_max(transpose([list(range(1000))])) == [999]
```

# 12. Define a function that takes two strings and decides whether they are anagrams. A string is an anagram of another string if its letters can be rearranged so that it equals the other string.¶

For example:

```
abc -- bac
aabb -- abab
```

Counter examples:

```
abc -- aabc
abab -- aaab
```

<!-- grade_id: cell-7ea849b6ffc4421a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-231f78d971a5d48b, points: 0, grade: true, solution: false -->
```python
assert(are_anagrams("", "") == True)
assert(are_anagrams("abc", "bac") == True)
assert(are_anagrams("aabb", "abab") == True)
assert(are_anagrams("abab", "aaab") == False)
```

# =================== EXTRA LEVEL ===================

# 13. Define a function that takes a string and returns its bigram frequencies as a dictionary.

Character bigrams are pairs of subsequent characters. For example word apple contains the following bigrams: ap, pp, pl, le.

Do not use the `collections` module.

<!-- grade_id: cell-2669edc3d29b28b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5a3764a1d0909393, points: 0, grade: true, solution: false -->
```python
assert get_char_bigrams("apple") == {'ap': 1, 'pp': 1, 'pl': 1, 'le': 1}
assert get_char_bigrams("apple apple") == {'ap': 2, 'pp': 2, 'pl': 2, 'le': 2, 'e ': 1, ' a': 1}
```

# 14. Define a function that takes two vectors represented as lists and returns their L2 distance.

Also called Euclidean distance.

<!-- grade_id: cell-849b6ea4c65f058e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4c9f06a8534cbefc, points: 0, grade: true, solution: false -->
```python
assert abs(l2_distance([1, 0], [1, 0]) - 0) < 1e-5
assert abs(l2_distance([1, 0], [0, 1]) - 2 ** 0.5) < 1e-5
```

# 15. Create a function that takes a matrix of size NxM and returns the pairwise L2 distance of each row pair.

For input $A^{N \times M}$, the output should be $D^{N \times N}$ where: $D_{i,j} = L_2 (A_i, A_j)$

<!-- grade_id: cell-644d3391fc995ee9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8bc3cf03b303c7c3, points: 0, grade: true, solution: false -->
```python
D = pairwise_row_distance([[1, 0, 2], [0, 2, 1]])
# should be a matrix
assert is_matrix(D) == True
assert len(D) == 2
assert len(D[0]) == 2
assert D[0][0] == 0
assert abs(D[0][1] - 2.449489742783178) < 1e-3
assert abs(D[1][0] - 2.449489742783178) < 1e-3
assert D[1][1] == 0

# should be symmetric
assert transpose(D) == D
```