# Introduction to Python and Natural Language Technologies

__Laboratory 03, Object oriented programming__

__Feb 24, 2021__

__SUBMISSION DEADLINE: Feb 28 23:59__

# 1.1  Define a class named `A` with a 'constructor' that takes a single parameter and stores it in an attribute named `value`. Add a `print_value` method to the class.

Instantiate the class and call the `print_value` method.

<!-- grade_id: cell-6c7ba38d59f0b533, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
a = A("abc")
a.print_value()
```

# 1.2 Redefine the class's `__init__` so that it can be instantiated without a parameter. If it is called without a parameter, value should be 42.

<!-- grade_id: cell-1fa4f4a04cc3b613, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
a = A("abc")
a.print_value()  # prints abc
a = A()
a.print_value()  # prints 42
```

# 1.3 Define a class named `B`, whose `__init__` takes two parameters and stores one in a public attribute and the other in a private attribute named `this_is_public` and `__this_is_private` respectively.

Check the class's __dict__ attribute and find out the mangled name of the private attribute.

<!-- grade_id: cell-f4c8b3c52d081013, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6480bb6183af4ff0, points: 0, grade: true, solution: false -->
```python
b = B(1, 2)
assert b.this_is_public == 1
try:
    b.__this_is_private
    print("This should not happen.")
except AttributeError:
    print("Failed to access private attribute, this is good :)")
```

# 2. Inheritance

# 2.1 Guess the output without running the cell.

```python
class A(object): pass
class B(A): pass
class C(A): pass
class D(B): pass

a = A()
b = B()
c = C()
d = D()

print(isinstance(a, object))
print(isinstance(b, object))
print(isinstance(a, B))
print(isinstance(b, A))
print(isinstance(d, A))
```

```python
print(issubclass(C, object))
print(issubclass(D, B))
print(issubclass(B, D))
print(issubclass(B, B))
```

# 2.2 Create a Cat, a Dog, a Fish and a Eagle class.

The animals have the following attributes:

1. cats, dogs and eagles can make a sound (this should be a make_sound function that prints the animals sound),
2. all animals have an age and a number_of_legs attribute,
3. cats and dogs have a fur_color attribute. They can be instantiated with a single color or a list or tuple of colors.

Use inheritance and avoid repeating code. Use default values in the constructors.

<!-- grade_id: cell-45b3470b12151035, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b584e9c19590038b, points: 0, grade: true, solution: false -->
```python
try:
    cat = Cat("Fluffy", age=3, fur="white")
    dog = Dog("Cherry", age=1, fur=("white", "brown", "black"))
    fish = Fish("King")
    eagle = Eagle("Bruce", age=2)
    animals = [cat, dog, fish, eagle]
except Exception as e:
    print(f'Could not initialize an animal: {e}')
```

Iterate over the list animals and call `make_sound` for each animal. Print either the sound the animal makes or "XY does not make a sound" if the animal does not make a sound (fish). This is an example of duck typing.

```python
for animal in animals:
    animal.make_sound()
```

# ====================== PASSING LEVEL ======================

# 3. `RationalNumber` class

Write a class that represents a rational number. A number is rational if it is can be expressed as the quotient of two integers (`p` and `q`). Define the operators seen in the tests below.

Make sure that `p` and `q` are always relative primes (you can use `math.gcd` or - even better - write your own implementation of the [Euclidean algorithm](https://en.wikipedia.org/wiki/Euclidean_algorithm)).

<!-- grade_id: cell-0e26530985fd1a9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-911e58f5299b4a3b, points: 0, grade: true, solution: false -->
```python
# Test basic operators
r = RationalNumber(43, 2)
assert r + r == RationalNumber(43)  # q = 1 in this case
assert r * 2 == r + r
r1 = RationalNumber(3, 2)
r2 = RationalNumber(4, 3)
assert r1 * r2 == RationalNumber(12, 6)
assert r1 / r2 == RationalNumber(9, 8)
assert r1 == RationalNumber(6, 4)
r4 = RationalNumber(2, 3)
r4.p = 20 * 55
r4.q = 55
assert r4 == RationalNumber(20)
```

You should modify the class definition `RationalNumber` above so that the tests below pass. You do not need to modify the rest of the notebook.

Make the class usable as a dictionary key.

<!-- grade_id: cell-272f9aac55adf7b1, points: 0, grade: true, solution: false -->
```python
# Test whether class is usable as dict key
r1 = RationalNumber(3)
r2 = RationalNumber(3, 1)
r3 = RationalNumber(3, 2)

d = {r1: 1, r2: 2, r3: 12}
assert(len(d) == 2)
```

`p` and `q` can only be integers. Raise a `RationalNumberValueError` if someone tries to set them to anything else.

<!-- grade_id: cell-879048de3710da83, points: 0, grade: true, solution: false -->
```python
# Test numerator and denominator validation
try:
    r1.p = 3.4
except ValueError:
    print("This should happen")
else:
    print("This shouldn't happen")
    
try:
    r1.q = 3.4
except ValueError:
    print("This should happen")
else:
    print("This shouldn't happen")
```

Rational numbers may be negative. Make sure that `q` is never negative.

<!-- grade_id: cell-fe941e6b2cdbd1e7, points: 0, grade: true, solution: false -->
```python
# Test negative number behaviour
r = RationalNumber(3, -2)
assert r.p == -3 and r.q == 2
assert abs(r) == 1.5
```

Add a `from_str` factory method which parses the following formats:

<!-- grade_id: cell-a366521c21fe9379, points: 0, grade: true, solution: false -->
```python
# Test from_str factory method
r1 = RationalNumber(-3, 2)
r2 = RationalNumber(12, 34)
assert RationalNumber.from_str("-3/2") == r1
assert RationalNumber.from_str("-6/4") == r1
assert RationalNumber.from_str("12 / 34") == r2
assert RationalNumber.from_str("36/    102") == r2
```

# ====================== EXTRA LEVEL ======================