# Introduction to Python and Natural Language Technologies

__Laboratory 4, Numpy__

__March 4, 2021__

__Gábor Borbély, Judit Ács__

```python
import numpy as np

import matplotlib.pyplot as plt
%matplotlib inline
```

# 1. Numpy basic exercises

Use vectorization and avoid for loops in all exercises. Implement each exercise as a single function.

## 1.1 Implement standardization for 2D arrays.

Standardization is defined as:

\begin{equation*}
X_{std} = \frac{X - \mu}{\sigma},
\end{equation*}

where $\mu$ is the mean of each row and $\rho$ is the standard deviation of each column.

<!-- grade_id: cell-053e17e70271858e, points: 0, grade: false, solution: true -->
```python
def standardize(X):
    # YOUR CODE HERE
    raise NotImplementedError()

X_std = standardize(np.array([[1, 2], [2, 3], [1, 2]]))
X_std
```

<!-- grade_id: cell-c63a4f94adc0d2c8, points: 0, grade: true, solution: false -->
```python
X_std = standardize(np.array([[1, 2], [2, 3], [1, 2]]))

answer = np.array(
    [[-0.70710678, -0.70710678],
     [ 1.41421356,  1.41421356],
     [-0.70710678, -0.70710678]])
assert np.allclose(X_std, answer)
```

## 1.2 Implement normalization for 2D arrays.

Normalization is defined as:

\begin{equation*}
X_{norm} = \frac{X - X_{min}}{X_{max} - X_{min}}
\end{equation*}

<!-- grade_id: cell-7aefa84284167b19, points: 0, grade: false, solution: true -->
```python
def normalize(X):
    # YOUR CODE HERE
    raise NotImplementedError()

X = np.arange(6).reshape(2, 3)
X[1, 2] = -5
print(X)
normalize(X)
```

<!-- grade_id: cell-71ecc8c7e90aa23f, points: 0, grade: true, solution: false -->
```python
answer = np.array(
    [[0., 0., 1.],
     [1., 1., 0.]]
)
X_norm = normalize(X)
assert np.allclose(X_norm, answer)
```

## 1.3 Implement the softmax function.

Softmax is defined as for each row:

$$
x_i \mapsto \frac{\exp(x_i)}{\sum_{j=1}^n \exp(x_j)},
$$

in other words, softmax maps every row into a probability distribution (all elements are nonnegative and their sum is 1).

<!-- grade_id: cell-3f51c2b903a73f7d, points: 0, grade: false, solution: true -->
```python
def softmax(X):
    # compute softmax along the last dimension
    # YOUR CODE HERE
    raise NotImplementedError()

X = np.arange(12).reshape(3, -1)
softmax(X)
```

<!-- grade_id: cell-5b49ba20bb8b7f7d, points: 0, grade: true, solution: false -->
```python
answer = np.array(
    [[0.0320586 , 0.08714432, 0.23688282, 0.64391426],
     [0.0320586 , 0.08714432, 0.23688282, 0.64391426],
     [0.0320586 , 0.08714432, 0.23688282, 0.64391426]]
) 
X = np.arange(12).reshape(3, -1)

assert np.allclose(softmax(X), answer)

X = np.array([[100, 100], [-100, -100]])
assert np.allclose(softmax(X), np.array([[0.5, 0.5], [0.5, 0.5]]))
```

<!-- grade_id: cell-3cfccb032defb8e7, points: 0, grade: true, solution: false -->
```python
X = np.array([[0, 1, 2],
              [2, 1, 2]])

answer = np.array([[0.09003057, 0.24472847, 0.66524096],
                   [0.4223188 , 0.1553624 , 0.4223188 ]])

assert np.allclose(softmax(X), answer)
```

# 2. Vectorization

Rewrite the following examples into vectorized solutions (no for loops and list comprehensions).

## 2.1 Row-wise Euclidean norm

Write a function which has one parameter, a 2D array and it returns a vector of row-wise Euclidean norms of the input. Use `numpy` operations and vectorization, avoid `for` loops. The solution below is a _bad_ solution.

```python
def rowwise_norm(X):
    def my_dot(x, y):
        result = 0.0
        for i in range(len(x)):
            result += x[i] * y[i]
        return result
    return np.array([np.sqrt(my_dot(x, x)) for x in X])

X = np.arange(5)[:, None]*np.ones((5, 3));
print(X)
print(rowwise_norm(X))
print(rowwise_norm([[1], [-1], [1], [-1]]))
```

Your vectorized solution goes here:

<!-- grade_id: cell-594581dff2c12a65, points: 0, grade: false, solution: true -->
```python
def rowwise_norm(X):
    # YOUR CODE HERE
    raise NotImplementedError()
    
X = np.arange(5)[:, None] * np.ones((5, 3));
print(X)
print(rowwise_norm(X))
```

<!-- grade_id: cell-f4dad882c53237a7, points: 0, grade: true, solution: false -->
```python
X = np.arange(5)[:, None] * np.ones((5, 3));
print(X)
print(rowwise_norm(X))

assert np.allclose(rowwise_norm(X), np.array([0., 1.73205081, 3.46410162, 5.19615242, 6.92820323]))
assert np.allclose(rowwise_norm([[1], [-1], [1], [-1]]), np.ones(4))
```

## 2.2 Chessboard

Write a function which has one parameter, a positive integer $n$, and returns an $n\times n$ array of $\pm1$ values like a chessboard: $M_{i,j} = (-1)^{i+j}$.

```python
def chessboard(n):
    return np.array([[(-1)**(i + j) for j in range(n)] for i in range(n)])

chessboard(5)
```

Your vectorized solution:

<!-- grade_id: cell-e0d3c46f473033f4, points: 0, grade: false, solution: true -->
```python
def chessboard(n):
    # YOUR CODE HERE
    raise NotImplementedError()

chessboard(3)
```

<!-- grade_id: cell-3b842f320efaf57a, points: 0, grade: true, solution: false -->
```python
board5 = np.array([[ 1, -1,  1, -1,  1],
                   [-1,  1, -1,  1, -1],
                   [ 1, -1,  1, -1,  1],
                   [-1,  1, -1,  1, -1],
                   [ 1, -1,  1, -1,  1]])
assert np.allclose(chessboard(5), board5)
assert np.allclose(chessboard(0), np.array([]))
```

# 3. Broadcast quiz

Do the following operations work and if so, what is the shape of the resulting array?

Can you broadcast these arrays such that they can be added? What will be the shape of the result?

Try to figure it out before evaluating the cells.

```python
np.ones(3) + np.ones((3, 3))
```

```python
np.ones(3) + np.ones((4, 3))
```

```python
np.ones(3) + np.ones((3, 4))
```

```python
np.ones(3)[:, None] + np.ones((3, 4))
```

```python
np.ones((1, 2, 3)) + np.ones((1, 3))[:, None, :]
```

```python
np.ones((1, 2, 3)) + np.ones((1, 3))
```

```python
np.ones((4, 20, 3)) + np.ones((5, 3))
```

# 4. Numpy advanced exercises

## 4.1 Blockmatrix

Write a function named __`blockmatrix`__ that produces the following (block) matrix:
$$
\left(\begin{array}{ccc|ccc}
 1 & \cdots & 1& 0 & \cdots & 0 \\
 \vdots & & \vdots  & \vdots & & \vdots \\
  1& \cdots & 1 & 0 & \cdots & 0 \\\hline
  0 & \cdots & 0 & 1 & \cdots & 1 \\
  \vdots & & \vdots & \vdots & & \vdots \\
  0 & \cdots & 0 & 1 & \cdots & 1
\end{array}\right)
$$
The function should have 2 positive integer parameters, the size of the first square block and the size of the last square block. The other two blocks should have the appropriate size (may be rectangle).
Both the first and last blocks are constant $1$ matrices.
Return the resulted matrix.

Use matrix initializers: `ones`, `zeros` and concatenation.

<!-- grade_id: cell-1bf312e2fb3e05d4, points: 0, grade: false, solution: true -->
```python
def blockmatrix(n, m):
    # YOUR CODE HERE
    raise NotImplementedError()

blockmatrix(2, 5)
```

<!-- grade_id: cell-3a6ef7655e12b536, points: 0, grade: true, solution: false -->
```python
block_mtx_3_2 = np.array([[1., 1., 1., 0., 0.],
                          [1., 1., 1., 0., 0.],
                          [1., 1., 1., 0., 0.],
                          [0., 0., 0., 1., 1.],
                          [0., 0., 0., 1., 1.]])
assert np.allclose(blockmatrix(3, 2), block_mtx_3_2)
assert np.allclose(blockmatrix(0, 0), np.array([]))
assert np.allclose(blockmatrix(1, 1), np.eye(2))
```

# ================ PASSING LEVEL ====================

## 4.2 Blockmatrix from arbitrary square matrices

Write a blockmatrix function that takes any number of square matrices and returns a blockmatrix with these matrices in the diagonal. You can use one for loop for this solution (list comprehensions count as a for loop).

<!-- grade_id: cell-5ae5fec188fd7419, points: 0, grade: false, solution: true -->
```python
def any_blockmatrix(*matrices):
    # YOUR CODE HERE
    raise NotImplementedError()

A = np.arange(9).reshape(3, 3)
B = np.diag((2, 3))
any_blockmatrix(A, B, A, B)
```

<!-- grade_id: cell-b3fc092f12ccd7cd, points: 0, grade: true, solution: false -->
```python
block_mtx = np.array([[0., 1., 2., 0., 0., 0., 0., 0., 0., 0.],
                      [3., 4., 5., 0., 0., 0., 0., 0., 0., 0.],
                      [6., 7., 8., 0., 0., 0., 0., 0., 0., 0.],
                      [0., 0., 0., 2., 0., 0., 0., 0., 0., 0.],
                      [0., 0., 0., 0., 3., 0., 0., 0., 0., 0.],
                      [0., 0., 0., 0., 0., 0., 1., 2., 0., 0.],
                      [0., 0., 0., 0., 0., 3., 4., 5., 0., 0.],
                      [0., 0., 0., 0., 0., 6., 7., 8., 0., 0.],
                      [0., 0., 0., 0., 0., 0., 0., 0., 2., 0.],
                      [0., 0., 0., 0., 0., 0., 0., 0., 0., 3.]])
A = np.arange(9).reshape(3, 3)
B = np.diag((2, 3))
assert np.allclose(any_blockmatrix(A, B, A, B), block_mtx)

block_mtx2 = np.array([[1., 1., 0., 0., 0.],
                       [1., 1., 0., 0., 0.],
                       [0., 0., 3., 3., 3.],
                       [0., 0., 3., 3., 3.],
                       [0., 0., 3., 3., 3.]])
assert np.allclose(any_blockmatrix(np.ones((2, 2)), 3 * np.ones((3, 3))), block_mtx2)
```

## 4.3 Derivative

Write a function which numerically derivates a $\mathbb{R}\mapsto\mathbb{R}$ function. Use the forward finite difference.

The input is a 1D array of function values, and optionally a 1D vector of abscissa values. If not provided then the abscissa values are unit steps.

The result is a 1D array with the length of one less than the input array.

Rewrite the following solution without for loops.

```python
def derivate(f, x=None):
    if x is None:
        x = np.arange(len(f))
    return np.array([(f[i+1] - f[i]) / (x[i+1] - x[i]) for i in range(len(x) - 1)])

derivate(np.arange(10)**2)
```

```python
x = np.arange(10)
plt.plot(x, x**2)
plt.plot(x[:-1], derivate(x**2, x))
```

Your vectorized solution:

<!-- grade_id: cell-3d9d5735780c3e94, points: 0, grade: false, solution: true -->
```python
def derivate(f, x=None):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a1fc0db52d9a48fb, points: 0, grade: true, solution: false -->
```python
x = np.arange(10)
plt.plot(x, x**2)
plt.plot(x[:-1], derivate(x**2, x))
```

## 4.4 Birthday problem

In probability theory, the birthday problem or birthday paradox concerns the probability that, in a set of n randomly chosen people, some pair of them will have the same birthday. By the pigeonhole principle, the probability reaches 100% when the number of people reaches 367 (since there are only 366 possible birthdays, including February 29). However, 99.9% probability is reached with just 70 people, and 50% probability with 23 people. These conclusions are based on the assumption that each day of the year (excluding February 29) is equally probable for a birthday. -- [Wikipedia](https://en.wikipedia.org/wiki/Birthday_problem)

Write a function that simulates this problem for variable $n$. Your function should take $n$ and an experiment count as its parameter and sample experiment count times and return the ratio of "birthday collisions" (how many times there were at least two birthdays on the same day). You can ignore leap years and assume that every year has 365 days.

Run it for different $n$ values with at least 1000 experiments.

<!-- grade_id: cell-0d94f23286e48f13, points: 0, grade: false, solution: true -->
```python
def simulate_birthday_problem(n, iter_no=1000):
    """Simulates the birthday problem for a group of N people
    Generates N birthdays with uniform probability and returns
    the proportion of iterations with birthday collisions.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

max_n = 50
fig, ax = plt.subplots(figsize=(16, 8))
ax.set_yticks(np.arange(0, 1.0, .1))
ax.grid(True)

# number of simulations to run
iter_nums = [20, 100, 1000]

for iter_no in iter_nums:
    x = []
    for n in range(2, max_n):
        x.append(simulate_birthday_problem(n, iter_no=iter_no))
    ax.plot(x)
    
ax.legend(iter_nums)

# notice that the plot becomes less noisy as we run more simulations
```

## 4.5 Horner's method

Implement the [Horner's method](https://en.wikipedia.org/wiki/Horner%27s_method#Description_of_the_algorithm) for evaluating polynomials. The first input is a 1D array of numbers, the coefficients, from the constant coefficient to the highest order coefficent. The second input is the variable $x$ to subsitute. The function should work for all type of variables: numbers, arrays; the output should be the same type array as the input, containing the elementwise polynomial values.

<!-- grade_id: cell-c4f302eefb8d1976, points: 0, grade: false, solution: true -->
```python
def horner(C, x):
    # YOUR CODE HERE
    raise NotImplementedError()

C = [2, 0, 1] # 2 + 0*x + 1*x^2
print(horner(C, 3)) # 2 + 3^2
print(horner(C, [3, 3]))
```

<!-- grade_id: cell-1c9876c20a694762, points: 0, grade: true, solution: false -->
```python
C = [2, 0, 1] # 2 + 0*x + 1*x^2
assert horner(C, 3) == 11
assert np.allclose(horner(C, [3, 3]), [11, 11])
```

# ================ EXTRA LEVEL ====================

# 5.1 K-means clustering

Implement the naive k-means algorithm for 2-dimensional data.

The naive version of k-means is explained [here](https://en.wikipedia.org/wiki/K-means_clustering#Standard_algorithm_(naive_k-means)).

K-means is sensitive to initialization. We fixed the random seeds so that the algorithm should successfully converge and find the cluster in the generated dataset. You are free to change the seed and find different clusters. Some seeds will result in empty clusters (and possibly numpy warnings, don't worry about those).

We first generate the data. You don't need to change this cell.

```python
from sklearn.datasets import make_blobs

np.random.seed(7)
sklearn_random_state = 0

# number of samples per cluster
num_samples_per_cluster = 200
# number of clusters
K = 3
X, y = make_blobs(n_samples=[num_samples_per_cluster] * K, centers=None, n_features=2, cluster_std=0.5,
                  random_state=sklearn_random_state)

# initial centroids
# with the above seeds this should be:
# array([[ 1.6905257 , -0.46593737],
#        [ 0.03282016,  0.40751628],
#        [-0.78892303,  0.00206557]]) 
centroids = np.random.standard_normal(size=(K, 2))
centroids
```

Implement k-means here:

<!-- grade_id: cell-e0d28ccfa8a8e22d, points: 0, grade: false, solution: true -->
```python
def run_kmeans(data, num_clusters, centroids, num_iter):
    # YOUR CODE HERE
    raise NotImplementedError()

cluster_assignments = run_kmeans(
    data=X, num_clusters=K, centroids=centroids, num_iter=10
)
```

Sanity checks:

<!-- grade_id: cell-8cc8c269789bd7d2, points: 0, grade: true, solution: false -->
```python
# cluster_assignment should be an integer for each sample
assert cluster_assignments.dtype == np.int64 or cluster_assignments.dtype == np.int32
assert cluster_assignments.shape == (3 * num_samples_per_cluster, )
# cluster_assignments should be in the range [0, K-1]
assert set(cluster_assignments) <= set(range(K))
```

Plot the original clusters and the output of k-means:

```python
fig, ax = plt.subplots(1, 2, figsize=(12, 4))
for class_value in range(K):
    row_ix = np.where(y == class_value)
    ax[0].scatter(X[row_ix, 0], X[row_ix, 1])
    
    row_ix = np.where(cluster_assignments == class_value)
    ax[1].scatter(X[row_ix, 0], X[row_ix, 1])
    
ax[0].set_title("Original")
ax[1].set_title("Cluster assignment")
```

# 5.2 Generate different datasets with `make_blobs` or `make_classification` and try running you k-means algorithm with different random seeds. Show a case where it's unsuccessful. 

```python

```

```python

```