# Introduction to Python and Natural Language Technologies

__Assignment 11, List comprehension, iteration, context managers, functional Python__

__April 28, 2021__


# 1. Comprehensions

Convert the following for loops into comprehensions:

```python
l = []
for i in range(-5, 10, 2):
    l.append(i-2)
l
```

<!-- grade_id: cell-fe81859f4057f7ff, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l = []
for i in range(100):
    if i % 10 == 4:
        l.append(i)
l
```

<!-- grade_id: cell-adadaf2661704ec1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l1 = [12, 1, 0, 13, -3, -4, 0, 2]
l2 = []

for e in l1:
    if e % 2 == 1:
        l2.append(e)
l2
```

<!-- grade_id: cell-01c992d67ca14968, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l1 = [12, 1, 0, 13, -3, -4, 0, 2]
l2 = []

for e in l1:
    if e % 2 == 1:
        l2.append(True)
    else:
        l2.append(False)
l2
```

<!-- grade_id: cell-106d47d0b13dd3b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l1 = [3, 5, 7, 11]
l2 = [2, 4, 6]

products = []

for x in l1:
    for y in l2:
        products.append(x*y)
products
```

<!-- grade_id: cell-4ead65faba295143, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
l1 = [3, 5, 7, 11, 13, 17, 19]
l2 = [2, 4, 6, 8, 10]

products = []

for x in l1:
    for y in l2:
        if (x + y) % 3 == 0:
            products.append(x*y)
products
```

<!-- grade_id: cell-c8caea611b1907b3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
fruits = ["apple", "plum", "pear", "avocado"]

mtx = []
for fruit in fruits:
    row = []
    for i, c in enumerate(fruit):
        row.append(c*(i+1))
    mtx.append(row)
    
mtx
```

<!-- grade_id: cell-b2e0df6824c6b028, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
text = "ababaacdsadb"

char_freqs = {}

for c in text:
    try:
        char_freqs[c] += 1
    except KeyError:
        char_freqs[c] = 1
        
char_freqs
```

<!-- grade_id: cell-9af4ab90dccea74a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
d1 = {"a": 1, "b": 3, "c": 2}
d2 = {"a": 2, "b": 1}

d3 = {}

for key in set(d1.keys()) | set(d2.keys()):
    max_val = max(d1.get(key, 0), d2.get(key, 0))
    d3[key] = max_val

d3
```

<!-- grade_id: cell-269ea11b9d01f4dd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 2. N-gram generators

## 2.1 Bigram generator
Write a generator function that takes a string and generates character bigrams from it:

<!-- grade_id: cell-e3f683ca2b1a38ff, points: 0, grade: false, solution: true -->
```python
from types import GeneratorType

def get_bigrams(text):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a9c433466ba1aab5, points: 0, grade: true, solution: false -->
```python
assert list(get_bigrams("abca")) == ["ab", "bc", "ca"]
assert list(get_bigrams("")) == []
assert isinstance(get_bigrams("a"), GeneratorType)
```

## 2.2 N-gram generator

Change the function so that it can generate n-grams. N should be a parameter of the function:

<!-- grade_id: cell-07f2eb9be62b41bf, points: 0, grade: false, solution: true -->
```python
def get_ngrams(text, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4117b15c31f462b3, points: 0, grade: true, solution: false -->
```python
assert list(get_ngrams("abca", 2)) == ["ab", "bc", "ca"]
assert list(get_ngrams("abca", 3)) == ["abc", "bca"]
assert list(get_ngrams("abca", 3)) == ["abc", "bca"]
assert isinstance(get_ngrams("a", 1), GeneratorType)
```

# 3. Generators with memory

Write a generator function that takes an iterator and a memory size `k` and yields a tuple of k elements from the iterator. You can base the solution on this one from Lecture 4:

```python
def iter_with_memory(orig_iter):
    prev = None
    for current in orig_iter:
        yield current, prev
        prev = current
```

<!-- grade_id: cell-213760247532dbf9, points: 0, grade: false, solution: true -->
```python
def iter_with_k_memory(orig_iter, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10d456c5035c4b76, points: 0, grade: true, solution: false -->
```python
l = list(range(3))
assert list(iter_with_k_memory(l, 3)) == [(0, None, None), (1, 0, None), (2, 1, 0)]
assert list(iter_with_k_memory(l, 1)) == [(0,), (1,), (2,)]
assert isinstance(iter_with_k_memory([], 1), GeneratorType)
```

# 4. Formatted printing with comprehensions

Given a list of lists of students, print them in the following format with one print invocation. You can find similar examples in Lecture 11.

```python
all_students = [
    ["Joe", "John", "Mary"],
    ["Tina", "Tony", "Jeff", "Béla"],
    ["Pete", "Dave"],
]
```

One class-per-line and print the size of the class too

~~~
class 1, size: 3, students: Joe, John, Mary
class 2, size: 2, students: Pete, Dave
~~~

<!-- grade_id: cell-cc231a351b5a051f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Sort the classes by size in increasing order

~~~
class 1, size: 2, students: Pete, Dave
class 2, size: 3, students: Joe, John, Mary
~~~

<!-- grade_id: cell-7b6657c1ef2e4231, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 6. Handling real text data

The following piece of code downloads a small sample of the Hungarian Webcorpus. We will work on this in later exercises.

The corpus contains a single word-per-line and sentence boundaries are denoted by empty lines.

The file has 4 columns separated by TABs:
1. original word
2. lemma (stemmed word)
3. morphological analysis
4. morphological analysis candidates.

Take a look at the file before continuing.

```python
import os
import urllib.request

fn = 'web2-4p-9-17'
zipname = fn + '.zip'

if not os.path.exists(zipname):
    print("Downloading the corpus")
    webcorp_url = "http://avalon.aut.bme.hu/~judit/resources/webcorp_parts/web2-4p-9-17.zip"
    u = urllib.request.URLopener()
    u.retrieve(webcorp_url, zipname)

if not os.path.exists(fn):
    from zipfile import ZipFile
    with ZipFile(zipname) as myzip:
        myzip.extractall()
```

## 6.1 Write a generator function that yields one sentence at a time as a list of tokens. Make sure that it yields the very last sentence of the file as well.

<!-- grade_id: cell-a510a2976c86dc7d, points: 0, grade: false, solution: true -->
```python
def read_sentences(filename):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-24bddf95f8bba368, points: 1, grade: true, solution: false -->
```python
sentence_iter = read_sentences(fn)

# The first sentence in 19 token long
sentence = next(sentence_iter)
assert(len(sentence) == 19)
# The second sentence in 7 token long
sentence = next(sentence_iter)
assert(len(sentence) == 7)

assert isinstance(sentence, list)

sentences = read_sentences(fn)

import types
assert isinstance(sentences, types.GeneratorType)

sentences = list(sentences)

assert(len(sentences) == 90764)
```

## 6.2 Write a generator function that yields one sentence at a time but skips short sentences. The length limit should be a parameter of the generator which defaults to 5.

<!-- grade_id: cell-afc44538de2e3868, points: 0, grade: false, solution: true -->
```python
def read_long_sentences(filename, min_length=5):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d1d5c221ba6496e1, points: 0, grade: true, solution: false -->
```python
sentences = read_long_sentences(fn)
assert isinstance(sentences, types.GeneratorType)

sentences = list(sentences)
assert len(sentences) == 85163

sentences = read_long_sentences(fn, 15)

sentences = list(sentences)
assert len(sentences) == 50059
```

# ===================== PASSING LEVEL ======================

# 7. Timer context manager

Create a `Timer` context manager that measures the running time of the `with` block. The context manager takes an optional name argument and prints the block's name at the end too. 

<!-- grade_id: cell-770fbefce0c11658, points: 0, grade: false, solution: true -->
```python
from datetime import datetime

class Timer:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-785c5443c7526939, points: 0, grade: true, solution: false -->
```python
# prints "slow code ran for F seconds
# F is the total_seconds the block took to finish (float)
with Timer("slow code"):
    s = sum(range(100000))
    
# prints "unnamed ran for F seconds
with Timer():
    s = sum(range(100000))
```

# 8. Binary search tree

Create a binary search tree for integers. Write tests for your solution as well.

Implement the following:
- iteration protocol for the tree. Traversal should be in-order (increasing order).
- `sum(tree)` - sum of all the elements
- `min(tree)`, `max(tree)` - smallest, largest element
- `len(tree)` - number of nodes

<!-- grade_id: cell-7e2a387ca8e66a85, points: 0, grade: false, solution: true -->
```python
class Tree:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6865e938aae8b739, points: 0, grade: true, solution: false -->
```python
t = Tree() 
assert len(t) == 0
t.insert(2)
assert 2 in t
assert len(t) == 1

try:
    t.insert(2.1)
except ValueError:
    pass
else:
    raise Exception("Should have raised a ValueError")

t.insert(-10)
t.insert(100)
assert len(t) == 3

# inserting the same value again
t.insert(2)
assert len(t) == 3
print(t)

# min, max and sum are supported for iterables
assert sum(t) == 92
assert min(t) == -10
assert max(t) == 100
```

# ===================== EXTRA LEVEL ======================

```python

```