# Introduction to Python and Natural Language Technologies

__Laboratory 12, Decorators__

__May 6, 2021__

__Judit Ács__


# 1. Function factories

## 1.1 Linear polynomial

Write a function factory that creates linear polynomials. A linear or first degree polynomial is a function that has two unbound parameters ($a$ and $b$) and bound variable ($x$):

\begin{equation*}
f(x) = ax + b.
\end{equation*}

Your function factory should return $f(x)$ for fixed $a$ and $b$ values.

<!-- grade_id: cell-65ef8ba699dc5597, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bbfd7e104190c1e5, points: 0, grade: true, solution: false -->
```python
f1 = create_linear_polynomial(1, 0)
f2 = create_linear_polynomial(2, -1)

assert f1(3) == 3 # 1*3+0 = 3
assert f2(5) == 9 # 2*5-1 = 9
```

## 1.2 Higher order polynomial

Create a polynomial factory for arbitrary degree polynomials:

\begin{equation*}
f(x) = a_n  x^n + a_{n-1} x^{n-1} + \dots + a_1 x + a_0.
\end{equation*}

An n-th order polynomial has $n+1$ coefficients including $a_0$.

<!-- grade_id: cell-0bf01c51186efebf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd9815b4e215cbcf, points: 0, grade: true, solution: false -->
```python
f1 = create_polynomial(1, 1, 1)
assert f1(2) == 7  # 1*x*x + 1*x + 1 = 7, if x = 2

f2 = create_polynomial(3)
assert f2(10) == 3 and f2(-100) == 3  # constant function

f3 = create_polynomial(3, 2, 1, -10)
assert f3(2) == 24  # 3*x*x*x + 2*x*x + 1*x - 10 = 24 + 8 + 2 - 10 = 24, if x=2
```

# 2. Decorators without parameters

## 2. 1 Timer decorator

Write a timer decorator that prints how many seconds the wrapped function takes to finish. Make sure that the original return value is returned by the wrapped function and function metadata is kept.

Create both the function (lowercase `timer`) and the class (uppercase `Timer`) version of the decorator.

### function version

<!-- grade_id: cell-6318811453fa0656, points: 0, grade: false, solution: true -->
```python
from functools import wraps
from datetime import datetime

def timer(func):
    # YOUR CODE HERE
    raise NotImplementedError()

@timer
def fast_func(x, y):
    """Fast function"""
    for _ in range(1000):
        pass
    print(x + y)

@timer
def slow_func():
    for _ in range(10000000):
        pass

@timer
def func_with_return(x):
    return 2 * x

fast_func(2, 3)
slow_func()
```

<!-- grade_id: cell-148b6c1b9ca60a46, points: 0, grade: true, solution: false -->
```python
# Function metadata should be kept intact.
assert fast_func.__doc__ == "Fast function"

# Check return value.
assert func_with_return(3) == 6
```

### class version

<!-- grade_id: cell-0d12c77dea414df3, points: 0, grade: false, solution: true -->
```python
class Timer:
    # YOUR CODE HERE
    raise NotImplementedError()
   
@Timer
def fast_func(x, y):
    """Fast function"""
    for _ in range(1000):
        pass
    print(x + y)
    
@Timer
def slow_func():
    for _ in range(10000000):
        pass
    
@Timer
def func_with_return(x):
    return 2 * x

fast_func(2, 3)
slow_func()
```

<!-- grade_id: cell-b4599ad8f4d1ba5c, points: 0, grade: true, solution: false -->
```python
# metadata should be kept intact
assert fast_func.__doc__ == "Fast function"

# check return value
assert func_with_return(3) == 6
```

## 2.2 `<html>` wrapper

Create a decorator that wraps a function's output between `<html>` and `</html>`. Write a function and a class version as well. You can assume that the wrapped function returns a string.

### function version

<!-- grade_id: cell-ff6836aac485f05d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

@html_wrapper
def greeter(name):
    return f"Hello {name}"

@html_wrapper
def greeter_fullname(firstname, lastname, order="western"):
    if order == "western":
        return f"Hello {firstname} {lastname}"
    elif order == "eastern":
        return f"Hello {lastname} {firstname}"
```

<!-- grade_id: cell-37ddab99b88ebbdf, points: 0, grade: true, solution: false -->
```python
assert greeter.__name__ == 'greeter'
assert greeter("John") == "<html>Hello John</html>"
assert greeter("Peter") == "<html>Hello Peter</html>"

assert greeter_fullname("John", "Smith") == "<html>Hello John Smith</html>"
assert greeter_fullname("John", "Smith", order="eastern") == "<html>Hello Smith John</html>"
```

### class version

<!-- grade_id: cell-c52e9118db15eb5c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

@HTMLWrapper
def greeter(name):
    return f"Hello {name}"

@html_wrapper
def greeter_fullname(firstname, lastname, order="western"):
    if order == "western":
        return f"Hello {firstname} {lastname}"
    elif order == "eastern":
        return f"Hello {lastname} {firstname}"
```

<!-- grade_id: cell-d668b5334370c987, points: 0, grade: true, solution: false -->
```python
assert greeter("John") == "<html>Hello John</html>"
assert greeter("Peter") == "<html>Hello Peter</html>"

assert greeter_fullname("John", "Smith") == "<html>Hello John Smith</html>"
assert greeter_fullname("John", "Smith", order="eastern") == "<html>Hello Smith John</html>"
```

# 3. Decorators with parameters

You only need to write a **function version** for the following decorators.

## 3.1 Create and arbitrary tag adder.

The decorator should take the `tag` as its only parameter and wrap the fuction's output between `<tag>` and `</tag>`.

<!-- grade_id: cell-83625773b580b5f9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

@tag_wrapper("h1")
def greeter(name):
    return f"Hello {name}"

@tag_wrapper("span")
def current_time():
    return f"{datetime.now()}"
```

<!-- grade_id: cell-5ea808db6f5f4a59, points: 0, grade: true, solution: false -->
```python
print(greeter("John"))
assert greeter.__name__ == 'greeter'
assert greeter("John") == "<h1>Hello John</h1>"

ct = current_time()
print(ct)

assert ct.startswith("<span>") and ct.endswith("</span>")
```

## 3.2 Parameter type checker

Create a decorator that takes a type as its parameter and checks every argument of the wrapped function including keyword arguments. If any of them are not instances of that type, it should raise a `TypeError`.

<!-- grade_id: cell-0812c2e6ab3f8445, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

@check_type(int)
def age_printer(age):
    print(f"I am {age} years old")
    
@check_type(int)
def double_age(age):
    return age * 2

@check_type(int)
def compare_age(age1, age2=0):
    return age1 > age2
    
age_printer(18)
```

<!-- grade_id: cell-9334061eca3ff160, points: 0, grade: true, solution: false -->
```python
assert age_printer.__name__ == 'age_printer'
# return value should not change
assert double_age(12) == 24

try:
    double_age("12")
except TypeError:  # a TypeError should be raised since the function's parameter is not an integer
    pass

assert compare_age(age2=12, age1=11) == False

try:
    compare_age(20, "abc")
except TypeError:
    pass
```

## 3.3 Only one exception type

Write a decorator that makes sure that a function only raises one type of error. If it would raise another exception, raise this type instead.

<!-- grade_id: cell-2c36fd8922a92285, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def increment_number(x):
    return x + 1  # would raise TypeError if not a number

try:
    increment_number("12")
except TypeError:
    print("TypeError raised")
    
@only_one_exception(ValueError)
def increment_number(x):
    return x + 1  # would raise TypeError if not a number

try:
    increment_number("12")
except ValueError:
    print("ValueError raised")
```

# ================= PASSING LEVEL ===================

## 3.4 Argument validator

Create a decorator that takes an arbitrary callable as its parameter and wraps function with one positional parameter. The decorator should 'validate' the parameter with its callable. The callable returns `False` if an argument is invalid, your decorator should raise a `ValueError` when it happens.

The decorator should raise a `TypeError` if the function takes anything other than a single positional argument.

<!-- grade_id: cell-783477140305dc11, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
        
def is_positive_int_convertible(n):
    try:
        n = int(n)  # raises ValueError if not possible
    except ValueError:
        return False
    return n > 0
        

@check_param(is_positive_int_convertible)
def age_printer(age):
    print(f"Your age is {int(age)}")
        
age_printer(12)
age_printer(12.0)
age_printer("12")
```

<!-- grade_id: cell-783477140305dc12, points: 0, grade: true, solution: false -->
```python
try:
    age_printer("abc")
except ValueError:
    print("ValueError raised, this should happen")
    
try:
    age_printer(12, 13)
except TypeError:
    print("TypeError raised, this should happen")
```

## 3.5 Check all arguments

<!-- grade_id: cell-6236e49e197c6605, points: 0, grade: false, solution: true -->
```python
import string
from functools import wraps


# YOUR CODE HERE
raise NotImplementedError()
        

def is_positive_int_convertible(n):
    try:
        n = int(n)  # raises ValueError if not possible
    except ValueError:
        return False
    return n > 0
        
    
def is_hungarian_name(name):
    if not name.istitle():
        return False
    non_ascii = set(list("áéíóöőúüű")) 
    letters = non_ascii | set(string.ascii_lowercase) | set(' ')
    name_letters = set(list(name.lower()))
    return name_letters < letters
        

@check_all_arguments(is_hungarian_name, is_positive_int_convertible, height=is_positive_int_convertible)
def serialize_hun_person_data(name, age, height=160):
    return f"Name: {name}, age: {age}, height: {height}"
```

<!-- grade_id: cell-c339018b03d78467, points: 0, grade: true, solution: false -->
```python
assert serialize_hun_person_data.__name__ == 'serialize_hun_person_data'
assert serialize_hun_person_data("Péter Kovács", "25", height=180) == "Name: Péter Kovács, age: 25, height: 180"

icelandic_name = "Aðalheiður"
try:
    serialize_hun_person_data(icelandic_name, "25", height=180)
except ValueError:
    pass
else:
    raise AssertionError("Expected a ValueError.")
```

### 3.5.1 Using the previous decorators define a serializer function that only serializes names in your native language or another language besides English if you are Hungarian. Add tests for valid and invalid inputs.

<!-- grade_id: cell-76bf99472768dde0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# ================= EXTRA LEVEL ===================

Some of these exercises are borrowed from these sources:

- https://www.python-course.eu/python3_decorators.php
- https://github.com/manahl/PythonTrainingExercises